# Actividad Piloto

## Iterara lo número del 0 al 99

## Instrucciones

Cree una función llamada `numeritos` que retorne en una lista los enteros del 0 al 99.

<!-- grade_id: solucion, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Caso 1: no existe la función.
try:
    numeritos
    assert type(numeritos) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada numeritos.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    numeritos()
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna una cadena
assert type(numeritos()) == list, f"Tu función debe retornar un valor de tipo {list.__name__}."

# Caso 4: respuesta explicita
assert len(numeritos()) == 100, "Tu función debe retornar una lista de longitud 100."

# Caso 5: la lista contiene solo números enteros
for i in numeritos():
    assert type(i) == int, f"La lista que retorna tu función contiene elementos que no son números enteros ({int.__name__})."

# Caso 6: el primer elemento no es el número 0
assert numeritos()[0] == 0, "El primer elemento de la lista que retorna tu función no es el número 0."

# Caso 7: cada elemento subsecuente es uno mas que el anterior
for i in range(1, len(numeritos())):
    assert numeritos()[i] - numeritos()[i-1] == 1, "Los elementos de la lista que retorna tu función no están bien ordenados o son incorrectos."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")

```

## Créditos

**Autor(es):** Alejandro Mantilla, Tobia Gasparoni

**Fecha última actualización:**: 11/08/2021